# Banco Mundial: préstamos a países en desarrollo

Detalles sobre la organización en: https://www.worldbank.org/en/about/what-we-do

### Disponible en Kaggle en:
https://www.kaggle.com/theworldbank/world-banks-major-contracts

El Banco Mundial es una organización internacional que se fundó en 1944 para reconstruir Europa después de la Segunda Guerra Mundial. Es uno de una variedad de organizaciones que ayudan a dar forma y equilibrar la economía mundial. Hoy, su función principal es combatir la pobreza ofreciendo asistencia para el desarrollo a países del tercer mundo de ingresos medios y bajos.
Según el propio Banco Mundial, la organización tiene dos objetivos muy específicos para 2030:
* Poner fin a la pobreza extrema disminuyendo el porcentaje de personas que viven con menos de `$`1.90 por día a no más del 3% 
* Promover la prosperidad compartida fomentando el crecimiento de los ingresos del 40% inferior para cada país.

### Variables y significado

1. As of Date: 6-Sep-17, date when the file was generated. 
2. Fiscal Year: año fiscal en el que se ha llevado a cabo el préstamo.
3. Region: área geográfica del país al que se le ha concedido el préstamo.
4. Borrower Country: país al que se le ha concedido el préstamo.
5. Borrower Country Code: código de dicho país.
6. Project ID: código interno del préstamo.
7. Project Name: nombre del proyecto que se llevará a cabo.
8. Procurement Type: área para la cual se utilizará el préstamo.
9. Procurement Category: ídem al anterior pero con una categorización más amplia.
10. Procurement Method: tipo de adjudicación.
11. Product Line: área interna a la que se le ha concedido el préstamo.
12. Major Sector: economic sector económico al que se destinará.
13. WB Contract Number: número interno de contrato.
14. Contract Description: descripción.
15. Contract Signing Date: fecha en la que el contrato se oficializó.
16. Supplier: proveedor principal que trabaja en el proyecto.
17. Supplier Country: país del proveedor principal.
18. Supplier Country Code: código del país del proveedor.
19. Total Contract Amount (USD): cantidad total en dólares.
20. Borrower Contract Reference Number: número de referencia del contrato.

**Nombre completo del alumno:**  

**INSTRUCCIONES**: en cada celda debes responder a la pregunta formulada, asegurándote de que el resultado queda guardado en la(s) variable(s) que por defecto vienen inicializadas a `None`. No se necesita usar variables intermedias, pero puedes hacerlo siempre que el resultado final del cálculo quede guardado exactamente en la variable que venía inicializada a None (debes reemplazar None por la secuencia de transformaciones necesarias, pero nunca cambiar el nombre de esa variable). **No olvides borrar la línea *raise NotImplementedError()* de cada celda cuando hayas completado la solución de esa celda y quieras probarla**.

Después de cada celda evaluable verás una celda con código. Ejecútala (no modifiques su código) y te dirá si tu solución es correcta o no. En caso de ser correcta, se ejecutará correctamente y no mostrará nada, pero si no lo es mostrará un error. Además de esas pruebas, se realizarán algunas más (ocultas) a la hora de puntuar el ejercicio, pero evaluar dicha celda es un indicador bastante fiable acerca de si realmente has implementado la solución correcta o no. Asegúrate de que, al menos, todas las celdas indican que el código es correcto antes de enviar el notebook terminado.

### Sobre el dataset anterior (Major_Contract_Awards.csv) se pide:

**(1 punto)** Ejercicio 1

* Leerlo tratando de que Spark infiera el tipo de dato de cada columna, y **cachearlo**. 
* Puesto que existen columnas que contienen una coma enmedio del valor, en esos casos los valores vienen entre comillas dobles. Spark ya contempla esta posibilidad y puede leerlas adecuadamente **si al leer le indicamos las siguientes opciones adicionales** además de las que ya sueles usar: `.option("quote", "\"").option("escape", "\"")`.
* Asegúrate de que las **filas que no tienen el formato correcto sean descartadas**, indicando también la opción `mode` con el valor `DROPMALFORMED` como vimos en clase.

```python
from pyspark.sql import SparkSession
from pyspark import SparkContext
from pyspark.sql.functions import *
from pyspark.sql.types import *
import pyspark.sql.functions as F
from pyspark.sql.types import IntegerType


spark = SparkSession\
    .builder\
    .appName("ModeloB")\
    .getOrCreate()
```

<!-- grade_id: read_csv, points: 0, grade: false, solution: true -->
```python
contractsDF = spark.read\
                 .option("header", "true")\
                 .option("inferSchema", "true")\
                 .option("quote", "\"")\
                 .option("escape", "\"")\
                 .option("mode", "DROPMALFORMED")\
                 .csv("ModeloB Major_Contract_Awards.csv").cache()
```

```python
contractsDF.count()
```

Output:
```
148515
```

<!-- grade_id: read_csv_test, points: 1, grade: true, solution: false -->
```python
from pyspark.sql.types import DoubleType
assert(contractsDF.count() == 148515)
```

**(1 punto)** Ejercicio 2

* La columna **Total Contract Amount (USD)** es en realidad numérica, pero todas las cantidades incluyen el signo `$` por lo que Spark la reconoce como string. Para corregir este comportamiento, vamos a eliminar el `$` de todas las filas utilizando la función `F.regexp_replace("Total Contract Amount (USD)", "\$", "")` donde `"\$"` es el string que queremos reemplazar (hay que escaparlo poniendo `\` delante porque sino el `$` se interpreta como un carácter especial), y siendo el nuevo string el string vacío, `""`. Esta función pertenece al paquete `pyspark.sql.functions`, por lo que ya funciona de manera distribuida, y devuelve como resultado un objeto columna transformado. 

* Aplica esta función dentro de la función `withColumn` para **reemplazar** la columna `Total Contract Amount (USD)` ya existente por la columna devuelta por `regexp_replace`. La manera de utilizarla es totalmente análoga a la utilización de, por ejemplo, la función `F.when` dentro de `withColumn`. 
* Aprovecha también para hacer un casting del objeto columna devuelto por regexp_replace, que es una columna de strings, a una columna de enteros: `F.regexp_replace(...).cast(...)`. Almacena el DF resultante en la variable `contractsDFenteros`, **cachéala** y utilízala a partir de este momento para trabajar en las celdas posteriores, salvo que la celda indique lo contrario.

<!-- grade_id: convert_dollars, points: 0, grade: false, solution: true -->
```python
# No olvides los imports que necesites...
# LÍNEAS EVALUABLES, NO RENOMBRAR LAS VARIABLES
contractsDFenteros = contractsDF.withColumn("Total Contract Amount (USD)", F.regexp_replace("Total Contract Amount (USD)", "\$", "").cast(IntegerType()))

contractsDFenteros.cache()
```

Output:
```
DataFrame[As of Date: string, Fiscal Year: string, Region: string, Borrower Country: string, Borrower Country Code: string, Project ID: string, Project Name: string, Procurement Type: string, Procurement Category: string, Procurement Method: string, Product line: string, Major Sector: string, WB Contract Number: int, Contract Description: string, Contract Signing Date: string, Supplier: string, Supplier Country: string, Supplier Country Code: string, Total Contract Amount (USD): int, Borrower Contract Reference Number: string]
```

<!-- grade_id: convert_dollars_tests, points: 1, grade: true, solution: false -->
```python
from pyspark.sql.types import IntegerType
assert(len(contractsDFenteros.columns) == len(contractsDF.columns))
assert(contractsDFenteros.count() == contractsDF.count())
assert(contractsDFenteros.schema["Total Contract Amount (USD)"].dataType == IntegerType())
```

**(1 punto)** Ejercicio 3

* Partiendo de `contractsDFenteros`, crear un nuevo DF donde la columna "Region" sea reemplazada por otra con mismo nombre, de tipo string en la que todos los valores de la columna original (LATIN AMERICA AND CARIBBEAN, SOUTH ASIA, OTHER ... etc) estén traducidos al español. Puedes elegir la traducción que más te guste, pero debe mantenerse el mismo número de categorías que ya había, que eran siete.
* El evaluador oculto comprobará que sigue habiendo el mismo número de ejemplos en cada categoría con el nuevo nombre, y que las categorías efectivamente se han traducido (ninguna se debe llamar igual que antes). Puedes cambiar AFRICA por África.

```python
contractsDFenteros.select((F.col("Region"))).distinct().show(20,False)
```

Output:
```
+----------------------------+
|Region                      |
+----------------------------+
|LATIN AMERICA AND CARIBBEAN |
|SOUTH ASIA                  |
|OTHER                       |
|AFRICA                      |
|MIDDLE EAST AND NORTH AFRICA|
|EAST ASIA AND PACIFIC       |
|EUROPE AND CENTRAL ASIA     |
+----------------------------+
```

<!-- grade_id: renombrar_region, points: 0, grade: false, solution: true -->
```python
# Voy a crear un diccionario que utilice la región en inglés como clave y la región en español como valor

paisesDict = {
  "LATIN AMERICA AND CARIBBEAN": "LATINOAMERICA Y CARIBE",
  "SOUTH ASIA": "SURASIA",
  "OTHER": "OTRO",
  "AFRICA": "África",
  "MIDDLE EAST AND NORTH AFRICA": "ORIENTE MEDIO Y NORAFRICA",
  "EAST ASIA AND PACIFIC": "ASIA ORIENTAL Y PACIFICO",
  "EUROPE AND CENTRAL ASIA": "EUROPA Y ASIA CENTRAL"
}

contractsTranslatedDF = contractsDFenteros.na.replace(paisesDict,"Region")

# contractsTranslatedDF = contractsDFenteros.na.replace(translation,"Region")

contractsTranslatedDF.select("Region").show(10)
```

Output:
```
C:\ProgramData\Anaconda3\lib\site-packages\pyspark\sql\dataframe.py:2213: UserWarning: to_replace is a dict and value is not None. value will be ignored.
  warnings.warn("to_replace is a dict and value is not None. value will be ignored.")
+--------------------+
|              Region|
+--------------------+
|              África|
|              África|
|              África|
|              África|
|ASIA ORIENTAL Y P...|
|ASIA ORIENTAL Y P...|
|ASIA ORIENTAL Y P...|
|ASIA ORIENTAL Y P...|
|             SURASIA|
|             SURASIA|
+--------------------+
only showing top 10 rows
```

<!-- grade_id: renombrar_region_test, points: 1, grade: true, solution: false -->
```python
assert(contractsTranslatedDF.select("Region").distinct().count() == 7)
g1 = contractsDF.groupBy("Region").count().withColumnRenamed("Region", "R1")
g2 = contractsTranslatedDF.groupBy("Region").agg(F.count("*").alias("c2"))
joinedDF = g1.join(g2, F.col("count") == g2.c2)
assert(joinedDF.count() == 7)
assert(joinedDF.where(F.col("count") == F.col("c2")).count() == 7)
assert(joinedDF.where(F.col("R1") == F.col("Region")).count() == 0)
```

**(1 punto)** Ejercicio 4

* Partiendo de `contractsTranslatedDF`, crear un nuevo DataFrame de **una sola fila** que contenga, **por este orden de columnas**, el **número** de categorías distintas existentes en cada una de las columnas `Procurement Type`, `Procurement Category` y `Procurement Method`. Pista: crear cada una de estas tres columnas al vuelo con `select`(). Renombrar cada columna de conteo para que se llame igual que la propia columna que estamos contando.

<!-- grade_id: numero_categorias, points: 0, grade: false, solution: true -->
```python
# LÍNEA EVALUABLE, NO RENOMBRAR VARIABLES
numeroCategoriasDF = contractsTranslatedDF.select(countDistinct(F.col("Procurement Type")).alias("Procurement Type")\
                                                 ,countDistinct(F.col("Procurement Category")).alias("Procurement Category")\
                                                 ,countDistinct(F.col("Procurement Method")).alias("Procurement Method"))
```

<!-- grade_id: numero_categorias_test, points: 1, grade: true, solution: false -->
```python
assert(len(numeroCategoriasDF.columns) == 3)
assert(numeroCategoriasDF.count() == 1)
categorias = numeroCategoriasDF.collect()[0]
assert(categorias["Procurement Type"] == 60)
assert(categorias["Procurement Category"] == 5)
assert(categorias["Procurement Method"] == 18)
```

**(3 puntos)** Ejercicio 5

* Partiendo de `contractsDFenteros` definido anteriormente, crear un **pipeline** formado por dos etapas: un indexador de la columna categórica `Procurement Method` y un discretizador (bucketizer) de la columna numérica que habíamos convertido a entero al principio, `Total Contract Amount (USD)`, de manera que sea convertida en una columna de números reales empezando en 0 y cuya parte decimal siempre sea 0.
* Para el indexador, si una vez entrenado le llegasen etiquetas que no ha visto antes, deberá eliminar esas filas (recordar la opción adecuada que vimos en clase). La columna de salida debe llamarse `ProcurementIndexed`.
* Para el discretizador, utilizar como puntos de corte los siguientes: (-Inf, 0, 100000, 200000, 300000, 400000, Inf). La columna de salida debe llamarse `TotalDiscretized`.
* Una vez creados ambos, componerlos para crear un Pipeline, y aplicarlo a `contractsDFenteros` para entrenar y a continuación también para predecir (es decir, transformarlo). El DF resultante de la transformación debe almacenarse en la variable `contractsTransformedDF`

<!-- grade_id: pipeline, points: 0, grade: false, solution: true -->
```python
# LÍNEA EVALUABLE, NO RENOMBRAR VARIABLES
# imports necesarios..........

from pyspark.ml.feature import Bucketizer
from pyspark.ml.feature import StringIndexer
from pyspark.ml import Pipeline

bucketizer = Bucketizer()
bucketizer.setSplits([-float("inf"), 0, 100000,200000,300000,400000, float("inf")])
bucketizer.setInputCol("Total Contract Amount (USD)")
bucketizer.setOutputCol("TotalDiscretized")

indexer = StringIndexer(inputCol="Procurement Method", outputCol="ProcurementIndexed",handleInvalid="skip")
pipeline = Pipeline(stages=[indexer,bucketizer])
pipelineModel = pipeline.fit(contractsDFenteros)
contractsTransformedDF = pipelineModel.transform(contractsDFenteros)

pandasContracts = contractsTransformedDF.toPandas()

pandasContracts
```

Output:
```
As of Date Fiscal Year                   Region  \
0       09/06/2017 12:00:00 AM        2004                   AFRICA   
1       09/06/2017 12:00:00 AM        2004                   AFRICA   
2       09/06/2017 12:00:00 AM        2004                   AFRICA   
3       09/06/2017 12:00:00 AM        2004                   AFRICA   
4       09/06/2017 12:00:00 AM        2004    EAST ASIA AND PACIFIC   
...                        ...         ...                      ...   
148510  09/06/2017 12:00:00 AM        2018                   AFRICA   
148511  09/06/2017 12:00:00 AM        2018                   AFRICA   
148512  09/06/2017 12:00:00 AM        2018                   AFRICA   
148513  09/06/2017 12:00:00 AM        2018  EUROPE AND CENTRAL ASIA   
148514  09/06/2017 12:00:00 AM        2018                   AFRICA   

       Borrower Country Borrower Country Code Project ID  \
0                  Mali                    ML    P041723   
1                  Mali                    ML  
```

<!-- grade_id: pipeline_test, points: 3, grade: true, solution: false -->
```python
from pyspark.sql import functions as F
assert("TotalDiscretized" in contractsTransformedDF.columns)
assert("ProcurementIndexed" in contractsTransformedDF.columns)
assert(len(contractsTransformedDF.columns) == len(contractsDFenteros.columns) + 2)
assert(bucketizer in pipeline.getStages())
assert(indexer in pipeline.getStages())
from pyspark.sql.types import DoubleType
assert(contractsTransformedDF.schema["TotalDiscretized"].dataType == DoubleType())
assert(contractsTransformedDF.schema["ProcurementIndexed"].dataType == DoubleType())
assert(bucketizer.getSplits() == [-float("Inf"), 0, 100000, 200000, 300000, 400000, float("Inf")])
assert(bucketizer.getInputCol() == "Total Contract Amount (USD)")
assert(bucketizer.getOutputCol() == "TotalDiscretized")
assert(indexer.getInputCol() == "Procurement Method")
assert(indexer.getOutputCol() == "ProcurementIndexed")
assert(indexer.getHandleInvalid() == "skip")
```

**(3 puntos)** Ejercicio 6

* Añadir una nueva columna al DF `contractsDFenteros` llamada `Total Pais` (sin tilde) que contenga en cada fila el **importe total prestado al país (Borrower Country Code) de esa fila (es decir, el total del país al que corresponde cada proyecto)**. El nuevo DF debe tener el mismo número de filas y una columna más. **NO DEBE UTILIZARSE JOIN sino funciones de ventana** con una ventana por país, que debe almacenarse en la variable `paisWindow`. El resultado sería equivalente a una agrupación por países y agregación de suma de importe, y después juntar cada total del país a cada celda (como si fuese un join por el país), pero **todo esto debe hacerse exclusivamente con ventanas y sin usar JOIN**.
* Una vez hecho esto, añade una segunda columna nueva llamada `Porcentaje Pais` (sin tilde) que contenga el **porcentaje** que ha supuesto cada proyecto sobre el total destinado a ese país (dicho total ha sido calculado en el punto anterior). El porcentaje no debe ir en tanto por 1 sino en tanto por 100.
* Añadir una tercera columna llamada `Media Pais` (sin tilde) que contenga en cada fila **el importe medio destinado a los proyectos del país al que corresponde el proyecto**. Debe utilizarse la misma ventana definida en el primer apartado, cambiando solo la función de agregación aplicada.
* Añadir una cuarta columna llamada `Diff Porcentaje` que sea la diferencia, medida en porcentaje, entre el importe destinado al proyecto y el importe medio de un proyecto en ese país. Debe calcularse operando con las columnas existentes, restando a la columna `Total Contract Amount (USD)` el importe de `Media Pais`, dividiendo entre esta última y multiplicando por 100, **sin utilizar** en ningún caso la función `when`. Una diferencia positiva indicará que ese proyecto ha recibido más fondos que la media de los proyectos de ese país, y una diferencia negativa indicará lo contrario.
* El DF resultante debe almacenarse en una variable `porcentajesDF` y debe tener el mismo número de filas que `contractsDFenteros`.

<!-- grade_id: window, points: 0, grade: false, solution: true -->
```python
# LÍNEA EVALUABLE, NO RENOMBRAR VARIABLES
# imports necesarios..........
from pyspark.sql.window import Window

# Crea la ventana para obtener la deuda total de cada país
paisWindow = Window.partitionBy("Borrower Country")

# Leí mal el ejercicio y pensé que tenía que calcular la suma total de todo el dinero en lugar de la suma total para cada país
# Como para mí esto ha sido una odisea y más formativo que calcular el total para cada país
# Lo voy a dejar aquí para futura referencia
# Conclusión: 'suma' y la columna "Porcentaje Pais Total" no se piden.
# Utiliza COLLECT() para obtener el valor total de la suma de todas las deudas para poder calcular el porcentaje!!
# He tenido que hacerlo así porque sum() sólo devuelve una fila así que da fallos muy raros!!!
suma = contractsDFenteros.select(sum("Total Contract Amount (USD)")).collect()[0][0]
# print(suma)

# Calcula la deuda total de cada país + El porcentaje de deuda sobre la total
porcentajesDF = contractsDFenteros\
                                .withColumn("Total Pais",F.sum("Total Contract Amount (USD)").over(paisWindow))\
                                .withColumn("Numero Proyectos",F.count("Borrower Country").over(paisWindow))\
                                .withColumn("Porcentaje Pais",F.col("Total Contract Amount (USD)")/(F.sum("Total Contract Amount (USD)").over(paisWindow))*100)\
                                .withColumn("Porcentaje Pais Total",F.col("Total Pais")/suma*100)\
                                .withColumn("Media Pais", F.col("Total Pais")/F.col("Numero Proyectos"))\
                                .withColumn("Diff Porcentaje",(F.col("Total Contract Amount (USD)")-F.col("Media Pais"))/F.col("Media Pais")*100)

# No te asustes por los porcentajes, pues sólo enseña los primeros 20 países!
porcentajesDF.select(F.col("Borrower Country"),F.col("Total Pais"),F.col("Porcentaje Pais"),F.col("Numero Proyectos"),F.col("Media Pais"),F.col("Diff Porcentaje")).show(5)

porcentajesDF=porcentajesDF.drop("Numero Proyectos")

# TODO: Calcular la columna 'Media Pais' y 'Diff Porcentaje'!
```

Output:
```
+----------------+----------+-------------------+----------------+-----------------+-------------------+
|Borrower Country|Total Pais|    Porcentaje Pais|Numero Proyectos|       Media Pais|    Diff Porcentaje|
+----------------+----------+-------------------+----------------+-----------------+-------------------+
|      South Asia|  49951856|0.09131192242386349|              83|601829.5903614458| -92.42111043881933|
|      South Asia|  49951856|0.09131192242386349|              83|601829.5903614458| -92.42111043881933|
|      South Asia|  49951856| 1.0141565110213322|              83|601829.5903614458|-15.825009585229427|
|      South Asia|  49951856|0.27518497010401377|              83|601829.5903614458| -77.15964748136686|
|      South Asia|  49951856| 0.4794016062185958|              83|601829.5903614458|-60.209666683856554|
+----------------+----------+-------------------+----------------+-----------------+-------------------+
only showing top 5 rows
```

<!-- grade_id: window_tests, points: 3, grade: true, solution: false -->
```python
assert("Total Pais" in porcentajesDF.columns)
assert("Porcentaje Pais" in porcentajesDF.columns)
r = porcentajesDF.where("`Project ID` = 'P069947'").head()
assert(r["Total Pais"] == 70485209)
assert(r["Porcentaje Pais"] - 25.985824912571374 < 0.001)
assert(r["Media Pais"] - 597332.279661017 < 0.001)
assert(r["Diff Porcentaje"] - 2966.3273396834217 < 0.001)
```